# Module 5 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

```python
import pandas as pd

import sys

from nose.tools import assert_equal, assert_almost_equal
from nose.tools import assert_equal, assert_is_instance, assert_is_not

import warnings
warnings.filterwarnings('ignore')

import json
import xml.etree.ElementTree as ET

import csv
```

# Problem 1: Reading a space delimited file

- Write a function called "read_file_space" that takes in a file name and uses the Python csv module to read it in row by row. Assume that the file is space delimited. Note that the file will return a nest list - a list where each element of the list is also a list.


<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
def read_file_space(data_file):
    '''
    A function which reads a space delimited file.
    
    
    Parameters
    -------
    data_file: a string, the name of the file to be read in
    
    
    Returns
    --------
    
    data: a list where each element of the list contains a row of the data file
    
    '''
    
    # YOUR CODE HERE
    
    return data
```

```python
data_file = './data/dow_jones_index.data'
data = read_file_space(data_file)
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(data,list,msg='Your result is not a list')
assert_equal(len(data),751,msg='Some rows of the dataset are missing. Make sure you arent skipping any rows and make sure you specified the right delimiter.')
assert_equal(data[0][0][0],'q',msg='The dataset was not read in correctly. Make sure you specified the right delimiter.')
```

# Problem 2: Get column names from a csv file.

- Write a function called "get_column_names" which uses the Python csv module to take in a comma separated file as input and returns the column names of the file as a Python list.


- You can assume that the column names (or equivalently, the header) are contained in the *first line* of the file which is being read.

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
def get_column_names(data_file):
    '''
    A function which reads in a csv file and returns the column names.
    
    
    Parameters
    -------
    data_file: a string, the name of the file to be read in
    
    
    Returns
    --------
    
    data: a list where each element of the list contains a row of the data file
    
    '''
    
    
    # YOUR CODE HERE
    
    return data[0]
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
data_file = './data/dow_jones_index.data'
column_names = get_column_names(data_file)

true_names = ['quarter','stock','date','open','high','low','close','volume','percent_change_price',
 'percent_change_volume_over_last_wk','previous_weeks_volume','next_weeks_open','next_weeks_close',
 'percent_change_next_weeks_price','days_to_next_dividend','percent_return_next_dividend']

assert_equal(true_names,column_names,msg='The column names are not correct. Make sure you specified the delimiter and'
             'that you took the correct row from the file.')

assert_is_instance(column_names,list,msg='Your function should return a list containing the column names.')
```

# Problem 3: Read in a JSON file.


- Write a function called "read_json" which reads in a .json file containing data.


- Use the Python json module to read in the file. This module has already been imported as json.


<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
def read_json(data_file):
    '''
    A function which reads in a JSON file.
    
    
    Parameters
    -------
    data_file: a string, the name of the file to be read in.
    
    
    Returns
    --------
    
    data: a Python list containing the read in data from the json file.
    
    '''
    
    
    # YOUR CODE HERE
    
    return data
```

```python
data_file = './data/dow_jones.json'
data = read_json(data_file)
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
data_file = './data/dow_jones.json'
data = read_json(data_file)

assert_equal(len(data),751,msg='Some rows of the dataset are missing.'
             'Make sure you arent skipping any rows.')
assert_is_instance(data,list,msg='Your function should return a lists.')
```

# Problem 4: Write a JSON file.


- Write a function called "read_json" which reads in a .json file containing data.


- Use the Python json module to read in the file. This module has already been imported as json.


<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
def write_json(file_name,data):
    '''
    A function which reads in a JSON file.
    
    
    Parameters
    -------
    data_file: a string, the name of the file to be read in.
    
    
    Returns
    --------
    
    data: a Python list containing the read in data from the json file.
    
    '''
    
    
    # YOUR CODE HERE
    
    return 0 
```

```python
#Write our json file
file_name = 'dowjones_local.json'
data = write_json(file_name,data)
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
#let's us our "read_json" function to see if we can read in the json file we saved
try:
    data = read_json(file_name)
except FileNotFoundError:
    print('File not found. It looks like the data was not saved correctly')

assert_equal(len(data),751,msg='Some rows of the dataset are missing.'
             'Make sure you arent skipping any rows when saving the file.')
```

```python
#Remove the json file that we saved to clear up space
import os
os.remove('dowjones_local.json')
print('File removed!')
```

# Problem 5: Working with XML formatted data.


- Write a function called "read_book_xml" which reads in a .xml dataset containing information about several books, and returns the name of the book in a Python list.


- Assume that the xml functionality in Python has been imported as "ET".


- Assume that within the .xml file, the name of the book is in an element called "book".


- Assume that within the "book" element, the title is the 2nd element.

<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
def read_book_xml(data_file):
    '''
    A function which reads in an xml containing author and book information and returns the list of book titles.
    
    
    Parameters
    -------
    data_file: a string, the name of the file to be read in.
    
    
    Returns
    --------
    
    titles: a Python list containing the titles of each book.
    
    '''
    
    
    # YOUR CODE HERE
    
    return titles
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
book_titles = read_book_xml('./data/books.xml')

assert_equal(book_titles, ["XML Developer's Guide", 'Midnight Rain', 'Maeve Ascendant', "Oberon's Legacy", 
                           'The Sundered Grail', 'Lover Birds', 'Splish Splash', 'Creepy Crawlies', 'Paradox Lost', 
                           'Microsoft .NET: The Programming Bible', 'MSXML3: A Comprehensive Guide', 
                           'Visual Studio 7: A Comprehensive Guide'])
```

-----

**&copy; 2018: Robert J. Brunner at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode